---
title: "Lesson 2: Study Design"
output: learnr::tutorial
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
library(learnr)
library(openintro)
library(tidyverse)

gradethis::gradethis_setup()
knitr::opts_chunk$set(echo = FALSE)

population2000 <- read_csv("../../../Datasets/gss2002.csv") %>% 
  mutate_if(is.character, factor)                           
```

## 2000 Presidential Election {data-progressive=TRUE}

As a reminder, here is some vocabulary that may be useful for today's activity:

* **population**: the group of items or events that you're interested in. Can be finite or infinite. May be hypothetical (e.g. all humans who have ever lived or will ever live).
* **sample**: the data you have. Usually a small portion of the population. Ideally, looks like a smaller version of the population and has similar properties.
* **census**: a sample of the whole population.
* **bias**: we say a sample is **biased** if it doesn't resemble the population.
* **simple random sampling**: a sample is chosen by randomly selecting cases from a population
* **stratified sampling**: the population is divided into groups called **strata**. The strata are chosen so that similar cases are grouped together, then a second sampling method, usually simple random sampling, is employed within each stratum.
* **cluster sampling**:  we break up the population into many groups, called clusters. Then we sample a fixed number of clusters and include all observations from each of those clusters in the sample.
* **multistage sampling**: like a cluster sample, but rather than keeping all observations in each cluster, we would collect a random sample within each selected cluster.

### Designing a poll

![](images/chads.gif)

From Wikipedia: The 2000 United States presidential election was the 54th quadrennial presidential election, held on Tuesday, November 7, 2000. Republican candidate George W. Bush, the governor of Texas and eldest son of the 41st president, George H. W. Bush, won the election, defeating incumbent Vice President Al Gore. It was the fourth of five American presidential elections, and the first since 1888, in which the winning candidate lost the popular vote, and is considered one of the closest elections in US history, with longstanding controversy surrounding the ultimate results.

For today, I want us to step in our time machines, go back to the year 2000 and pretend that we're pollsters trying to predict the outcome of the election. 

You've been divided into "polling firms" (i.e. groups). The first thing I want you to do with your group is decide on a name for your polling firm. 

What is the population of interest in this scenario? Is it feasible to collect a *census* of the population?

```{r population, exercise=TRUE}
# Answer:
```

Take about 5 minutes to brainstorm how you would try and collect a sample to try an predict who would win the 2000 election. Try an use some vocabulary from above. Important things to consider:
* What information (i.e. variables) would you want to collect about each individual?
* What type of sampling would you use?
* How much labor and expense would it take to collect this sample?
* How will you try to eliminate bias in your sample?
Use the console below to write down your answer and then copy it into padlet:

```{r brainstorm, exercise=TRUE}
# Answer:
```

### Working with data

![One of George W. Bush's dog paintings](images/gwbdog.jpeg){width=50%}

I've assembled a dataset, stored as `population2000` which has 100,000 observations and 5 variables. For now, let's pretend that this dataset contains the true underlying population (i.e. all voters in the 2000 election) and our goal is the estimate the proportion of voters who will vote for George "Dubya" Bush and what proportion will vote for Al Gore. The variables in the dataset are as follows:

* `age`: the voter's age bucketed into equally sized levels `[18,36]`, `(36,53]`, `(53,89]`
<!-- * `race`: the voter's race with levels `white`, `black of african american`, `american indian or alska native`, `asian or pacific islander`, `hispanic`, `other` -->
<!-- * *degree*: the highest degree obtained by the voter with levels: -->
<!--   * `lt high school`: some high school -->
<!--   * `high school`: completed high school -->
<!--   * `junior college`: completed junior college -->
<!--   * `bachelor`: completed a bachelors degree -->
<!--   * `graduate`: completed a graduate degree -->
* `provider`: the voter's phone provider with levels `Sprint`, `Comcast`, `Verizon`, `AT&T`, `Time Warner`
* `pres00`: who they plan to vote for in the 2002 election. The variable has levels `gore`, `bush`, or `nader`

### Population Proportions

Let's first compute the true population proportion of voter's who plan to vote for Bush and Gore. To do that you must:
1. Pipe `population2000` into the `count` function, making sure to count the variable `pres00`. This will create a new column `n` which contains the counts for each candidate.
2. To get the proportions, pipe the output from the previous step into a `mutate` function and specify `prop = n/sum(n)` which will compute the proportion for each level.

```{r bush-popprop, exercise=TRUE}

```

```{r bush-popprop-hint-1}
____ %>%
  ____(____) %>%
  ____(____)
```

```{r bush-popprop-hint-2}
____ %>%
  count(____) %>%
  mutate(____)
```

```{r bush-popprop-hint-3}
population2000 %>%
  count(pres00) %>%
  mutate(prop = n / sum(n))
```

```{r bush-popprop-check}
grade_result(
  pass_if(~identical(.result, population2000 %>%
  count(pres00) %>%
  mutate(prop = n / sum(n))))
)
```

### Simple Random Sample

Now suppose that you only have enough funding to collect a sample of 300 voters. Take a simple random sample of 300 voters from `population2000`, and call it `pop2000_srs`:

```{r srs-setup}
set.seed(1989)
```

```{r srs, exercise=TRUE}

```

```{r srs-hint-1}
# You may find the function `slice_sample` useful.
```

```{r srs-hint-2}
pop2000_srs <- ____ %>%
  ____(n = ____)
```

```{r srs-hint-3}
pop2000_srs <- ____ %>%
  slice_sample(n = ____)
```

```{r srs-hint-4}
pop2000_srs <- population2000 %>%
  slice_sample(n = 300)
```

```{r srs-check}
set.seed(1989)
grade_result(
  pass_if(~identical(pop2000_srs, population2000 %>%
  slice_sample(n = 300)))
)
```

Now, using similar code as above, compute the sample proportion of voter's who plan to vote for Bush and Gore. To do that you must:
1. Pipe `pop2000_srs` into the `count` function, making sure to count the variable `pres00`. This will create a new column `n` which contains the counts for each candidate.
2. To get the proportions, pipe the output from the previous step into a `mutate` function and specify `prop = n/sum(n)` which will compute the proportion for each level.

```{r srs-gen}
set.seed(1989)
pop2000_srs <- population2000 %>%
  slice_sample(n = 300)
```

```{r srs-prop, exercise=TRUE, exercise.setup="srs-gen"}

```

```{r srs-prop-hint-1}
____ %>%
  ____(____) %>%
  ____(____)
```

```{r srs-prop-hint-2}
____ %>%
  count(____) %>%
  mutate(____)
```

```{r srs-prop-hint-3}
pop2000_srs %>%
  count(pres00) %>%
  mutate(prop = n / sum(n))
```

```{r srs-prop-check}
grade_result(
  pass_if(~identical(.result, pop2000_srs %>%
  count(pres00) %>%
  mutate(prop = n / sum(n))))
)
  
```

### Using phone providers

The way that your firm actually conducts your poll is by paying different phone providers for a list of numbers for you do dial. Suppose that you don't have enough money to pay all 5 telecom companies so you decide to randomly select 2 providers and conduct a simple random sample of size 150 in each. Assume that the properties of the population of voters that use each provider are very similar.

```{r ms-q}
question_radio(
  "What type of sample is this?",
  answer("Simple Random Sample"),
  answer("Stratified Sample"),
  answer("Cluster Sample", message="Very close! This would be a cluster sample if we were to survey all of the voters for the selected providers."),
  answer("Multistage Sample", correct=TRUE, message="Correct! This is a multistage sample since we first select a sample of providers and then a simple random sample of voters belonging to each provider."),
               allow_retry = TRUE
)
```

### Multistage Sampling

We have made the assumption that those using each provider are very similar. Below, brainstorm some reasons why this may not be true and then post those reasons in the padlet.

```{r not-similar, exercise=TRUE}
# Brainstorm: 
```

The code below will randomly select a sample of 2 providers. The `level` function get's the levels of the `provider` variable and the `sample` function takes a random sample of the input vector. Feel free to use `?` to look up the help files on these function. Remember that you can select a column by name by putting `$` between the name of the data frame and the name of the column.

```{r select-providers-setup}
set.seed(2000)
```

```{r select-providers, exercise=TRUE}
sample(levels(population2000$provider), size=2)
```

As we can see we're surveying customers of Verizon and Time Warner. In the console, the pre-populated code generates a data frame called `pop2000_verizon_tw` which contains only customers of Verizon and Time Warner. Note that `==` tests whether two things are equal and `|` stands for a [logical "or"](https://en.wikipedia.org/wiki/Logical_disjunction). Create a data frame containing counts of Bush and Gore voters within each cluster. To do this:
1. Collect a simple random sample of size 150 from each by piping `pop2000_verizon_tw` into the function `group_by`. Make sure to group by provider.
2. Pipe the result into the `slice_sample` function. Make sure to specify the sample size.
3. Pipe the result into `ungroup`. We haven't learned this function yet. We do this because so that when we eventually compute the proportion of voters who plan on voting for each candidate, it won't aggregate by `provider`.
3. Store the result in a data frame called `pop2000_provider_srs`.

```{r providers-srs-setup}
set.seed(5001)
```

```{r providers-srs, exercise=TRUE}
pop2000_verizon_tw <- population2000 %>%
  filter(provider == "Verizon" | provider == "Time Warner")

# Add code below:


# Add code above. The line below will print out your result
glimpse(pop2000_provider_srs)
```

```{r providers-srs-hint-1}
pop2000_provider_srs <- _____ %>%
  _____(_____) %>%
  _____(_____) %>%
  _____
```

```{r providers-srs-hint-2}
pop2000_provider_srs <- pop2000_verizon_tw %>%
  group_by(_____) %>%
  slice_sample(_____) %>%
  ungroup()
```

```{r providers-srs-hint-3}
pop2000_provider_srs <- pop2000_verizon_tw %>%
  group_by(provider) %>%
  slice_sample(n=150) %>%
  ungroup()
```

```{r providers-srs-check}
set.seed(5001)
grade_result(
  pass_if(~identical(pop2000_provider_srs, pop2000_verizon_tw %>%
  group_by(provider) %>%
  slice_sample(n=150) %>%
  ungroup()))
)
```

Compute the proportion of voters in `pop2000_provider_srs` who plan on voting for Gore and Bush in the console below:

```{r providers-prop-setup}
set.seed(5001)

pop2000_verizon_tw <- population2000 %>%
  filter(provider == "Verizon" | provider == "Time Warner")

pop2000_provider_srs <- pop2000_verizon_tw %>%
  group_by(provider) %>%
  slice_sample(n=150) %>%
  ungroup()
```

```{r providers-prop, exercise=TRUE}

```

```{r providers-prop-hint-1}
_____ %>%
  count(_____) %>%
  mutate(_____)
```

```{r providers-prop-hint-2}
pop2000_provider_srs %>%
  count(pres00) %>%
  mutate(prop = n / sum(n))
```

```{r providers-prop-check}
grade_result(
  pass_if(~identical(.result, pop2000_provider_srs %>%
  count(pres00) %>%
  mutate(prop = n / sum(n))))
)
```

Would you expect this sampling scheme to be more or less accurate than a simple random sample? Why? Enter your answer below and then copy that into the padlet:

```{r multi-accurate, exercise=TRUE}
# Answer:
```

### Using age

Your polling firm as done some research and has found that people's voting behaviors greatly vary based on their age. As a result, when you create your sample, you want to make sure that the sample you collect accurately represents the distribution of ages in the population. The variable `age` contains three age buckets with a (roughly) equal number of voters. To assemble your sample, you plan to select 100 voters from each bucket.

```{r strat-q}
question_radio(
  "What type of sample is this?",
  answer("Simple Random Sample"),
  answer("Stratified Sample", correct=TRUE),
  answer("Cluster Sample"),
  answer("Multistage Sample"),
               allow_retry = TRUE
)
```

### Stratified Sampling

In the console below assemble a sample of 300 voters by selecting 100 from each age bucket. You can reuse some of the code above when we generated a sample of 150 from two different telecome providers. Call the new data frame `pop2000_strat`:

```{r strat-gen-setup}
set.seed(2020)
```

```{r strat-gen, exercise=TRUE}
pop2000_strat <-
```

```{r strat-gen-hint-1}
# 1. Pipe population2000 into a group_by function (make sure to group by age)
# 2. Pipe this result into a slice_sample function (make sure to select the correct sample size)
# 3. Pipe this result into an ungroup function
```

```{r strat-gen-hint-2}
pop2000_strat <- _____ %>%
  _____(_____) %>%
  _____(_____) %>%
  _____
```

```{r strat-gen-hint-3}
pop2000_strat <- population2000 %>%
  group_by(_____) %>%
  slice_sample(_____) %>%
  ungroup()
```

```{r strat-gen-hint-4}
pop2000_strat <- population2000 %>%
  group_by(age) %>%
  slice_sample(n = 100) %>%
  ungroup()
```

```{r strat-gen-check}
set.seed(2020)

grade_result(
  pass_if(~identical(pop2000_strat, population2000 %>%
  group_by(age) %>%
  slice_sample(n = 100) %>%
  ungroup()))
)
```

Finally compute the proportion of voters who plan on voting for each candidate in this new stratified sample:

```{r strat-prop-setup}
set.seed(2020)

pop2000_strat <- population2000 %>%
  group_by(age) %>%
  slice_sample(n = 100) %>%
  ungroup()
```

```{r strat-prop, exercise=TRUE}

```

```{r strat-prop-hint-1}
_____ %>%
  _____(_____) %>%
  _____(_____)
```

```{r strat-prop-hint-2}
pop2000_strat %>%
  count(_____) %>%
  mutate(_____)
```

```{r strat-prop-hint-3}
pop2000_strat %>%
  count(pres00) %>%
  mutate(prop = n / sum(n))
```

```{r strat-prop-check}
grade_result(
  pass_if(~identical(.result, pop2000_strat %>%
  count(pres00) %>%
  mutate(prop = n / sum(n))))
)
```

### Using race

Your polling firm as done further research and has found that people's voting behaviors vary even more based on their race. The problem is that there aren't an equal number of voters of each race. Below brainstorm how you might generate a sample that is representative of the population and ways you might go about creating an unbiased estime of the proportion of voters who vote for each cadidate. Please then post this in the padlet.

```{r race-brainstorm, exercise=TRUE}
# Brainstorm: 
```


## Sample and Population Questions

### Sleeping in college

A recent article in a college newspaper stated that college students get an average of 5.5 hours of sleep each night. A student who was skeptical about this value decided to conduct a survey by randomly sampling 25 students. On average, the sampled students slept 6.25 hours per night. Identify which value represents the sample mean and which value represents the claimed population mean.

```{r sleeping-in-college, echo=FALSE}
quiz(
  caption="Sleeping in college",
  question_numeric("Sample mean",
                   answer(6.25, correct = TRUE),
                   allow_retry = TRUE),
  question_numeric("Claimed Population Mean",
                   answer(5.5, correct = TRUE),
                   allow_retry = TRUE)
)
```

### Cheaters, scope of inference

Researchers studying the relationship between honesty, age and self-control conducted an experiment on 160 children between the ages of 5 and 15. The researchers asked each child to toss a fair coin in private and to record the outcome (white or black) on a paper sheet and said they would only reward children who report white. Half the students were explicitly told not to cheat, and the others were not given any explicit instructions. Differences were observed in the cheating rates in the instruction and no instruction groups, as well as some differences across children’s characteristics within each group. (Bucciol and Piovesan 2011)

```{r cheaters, echo=FALSE}
quiz(
  caption="Cheaters",
  question_radio("Identify the population of interest in this study.",
                   answer("All human beings"),
                   answer("All children"),
                   answer("All children between the ages of 5 and 15", correct = TRUE),
                   answer("160 children between the ages of 5 and 15"),
                   allow_retry = TRUE),
  question_radio("Identify the sample in this study.",
                   answer("All human beings"),
                   answer("All children"),
                   answer("All children between the ages of 5 and 15"),
                   answer("160 children between the ages of 5 and 15", correct = TRUE),
                   allow_retry = TRUE)
)
```

Comment on whether the results of the study can be generalized to the population, and if the findings of the study can be used to establish causal relationships.

```{r cheater-generalization, exercise=TRUE}
# Answer: 
```

```{r cheater-generalization-solution}
# Answer: If the children in this sample, who are likely not randomly sampled, 
# can be considered to be representative of all children between the ages of 5 
# and 15, then the results are generalizable to the population defined above. 
# Additionally, since the study is experimental, the findings can be used to 
# establish causal relationships.
```

### Cats on YouTube

Suppose you want to estimate the percentage of videos on YouTube that are cat videos. It is impossible for you to watch all videos on YouTube, so you use a random video picker to select 1000 videos for you. You find that 2% of these videos are cat videos. Determine which of the following is an observation, a variable, a sample statistic, or a population parameter.

```{r cats, echo=FALSE}
quiz(
  caption="Cats on Youtube",
  question_radio("Percentage of all videos on YouTube that are cat videos.",
                 answer("Observation"),
                 answer("Variable"),
                 answer("Sample Statistic"),
                 answer("Population Parameter", correct = TRUE),
                 allow_retry = TRUE),
  question_radio("2%.",
                 answer("Observation"),
                 answer("Variable"),
                 answer("Sample Statistic", correct = TRUE),
                 answer("Population Parameter"),
                 allow_retry = TRUE),
  question_radio("A video in your sample.",
                 answer("Observation", correct = TRUE),
                 answer("Variable"),
                 answer("Sample Statistic"),
                 answer("Population Parameter"),
                 allow_retry = TRUE),
  question_radio("Whether a video is a cat video.",
                 answer("Observation"),
                 answer("Variable", correct = TRUE),
                 answer("Sample Statistic"),
                 answer("Population Parameter"),
                 allow_retry = TRUE)
)
```

## Types of Study and Sampling Questions

### Course satisfaction across sections.

A large college class has 160 students. All 160 students attend the lectures together, but the students are divided into 4 groups, each of 40 students, for lab sections administered by different teaching assistants. The professor wants to conduct a survey about how satisfied the students are with the course, and he believes that the lab section a student is in might affect the student’s overall satisfaction with the course.

```{r course-satisfaction}
quiz(
  caption = "Course satisfaction across sections",
  question_radio("What type of study is this?",
                 answer("Experiment"),
                 answer("Observational Study", correct = TRUE),
                 allow_retry = TRUE),
  question_radio("What sampling strategy would be best for carrying out this study.",
                 answer("Simple random sampling", message = "Not quite! With SRS we can't ensure that all section are represented."),
                 answer("Stratified sampling where each section is a strata", correct = TRUE, message = "The professor suspects students in a given section may have similar feelings about a course. To ensure each section is reasonably represented, she may choose to randomly select a fixed number of students, say 10, from each section for a total sample size of 40 students. Since a random sample of fixed size was taken within each section in this scenario, this represents a stratified sampling."),
                 answer("Cluster sampling where each section is a cluster", message = "Sorry! If we suspect that the section impacts performance then the clusters wouldn't resemble one another."),
                  answer("Multistage sampling where each section is a cluster", message = "Sorry! If we suspect that the section impacts performance then the clusters wouldn't resemble one another."),
                 allow_retry = TRUE)
)
```


### Stressed out

A study that surveyed a random sample of otherwise healthy high school students found that they are more likely to get muscle cramps when they are stressed. The study also noted that students drink more coffee and sleep less when they are stressed.

```{r stressed, echo=FALSE}
quiz(
  caption="Stressed out",
  question_radio("What type of study is this?",
                 answer("Experiment"),
                 answer("Observational Study", correct = TRUE),
                 allow_retry = TRUE),
  question_radio("Can this study be used to conclude a causal relationship between increased stress and muscle cramps?",
                 answer("Yes", message="Since the study is observational we cannot infer causation."),
                 answer("No", correct=TRUE, message="Since the study is observational we cannot infer causation."),
                 allow_retry = TRUE),
  question_checkbox("Which of the following are possible confounding variables that might explain the observed relationship between increased stress and muscle cramps?",
                    answer("Stress"),
                    answer("Muscle Cramps"),
                    answer("Caffeine", correct = TRUE),
                    answer("Lack of sleep", correct = TRUE),
                    allow_retry = TRUE,
                    random_answer_order = TRUE)
  
)
```

### Sampling strategies. 

A statistics student who is curious about the relationship between the amount of time students spend on social networking sites and their performance at school decides to conduct a survey. Various research strategies for collecting data are described below. In each, name the sampling method proposed and any bias you might expect.

They randomly sample 40 students from the study’s population, give them the survey, ask them to fill it out and bring it back the next day.

```{r sampling1, exercise=TRUE}
# Answer:
```

```{r sampling1-solution}
# Answer: Simple random sample. Non-response bias, if only those people who have strong opinions about the survey responds their sample may not be representative of the population.

```


They give out the survey only to their friends, making sure each one of them fills out the survey.

```{r sampling2, exercise=TRUE}
# Answer:
```

```{r sampling2-solution}
# Answer: Convenience sample. Their sample may not be representative of the population since it consists only of their friends. This is called "under coverage bias". It is also possible that the study will have non-response bias if some choose to not bring back the survey.
```

They randomly sample 5 classes and asks a random sample of students from those classes to fill out the survey.

```{r sampling3, exercise=TRUE}
# Answer:
```

```{r sampling3-solution}
# Answer: Multi-stage sampling. If the classes are similar to each other with respect to student composition this approach should not introduce bias, other than potential non-response bias.
```




### Evaluate sampling methods. 

A university wants to determine what fraction of its undergraduate student body support a new $25 annual fee to improve the student union. 

```{r evaluate-sampling-method, echo=FALSE}
question_checkbox("Select each method which is reasonable.",
                  answer("Survey a simple random sample of 500 students.", message = "Simple random sampling is okay. In fact, it’s rare for simple random sampling to not be a reasonable sampling method!", correct=TRUE),
                  answer("Stratify students by their field of study, then sample 10% of students from each stratum.", message = "The student opinions may vary by field of study, so the stratifying by this variable makes sense and would be reasonable.", correct=TRUE),
                  answer("Cluster students by their ages (e.g., 18 years old in one cluster, 19 years old in one cluster, etc.), then randomly sample three clusters and survey all students in those clusters.", message="Students of similar ages are probably going to have more similar opinions, and we want clusters to be diverse with respect to the outcome of interest, so this would not be a good approach. (Additional thought: the clusters in this case may also have very different numbers of people, which can also create unexpected sample sizes.)", correct=FALSE),
                  allow_retry = TRUE
)
```



## Haters are gonna hate, study confirms. 

A study published in the Journal of Personality and Social Psychology asked a group of 200 randomly sampled participants recruited online using Amazon’s Mechanical Turk to evaluate how they felt about various subjects, such as camping, health care, architecture, taxidermy, crossword puzzles, and Japan in order to measure their attitude towards mostly independent stimuli. Then, they presented the participants with information about a new product: a microwave oven. This microwave oven does not exist, but the participants didn’t know this, and were given three positive and three negative fake reviews. People who reacted positively to the subjects on the dispositional attitude measurement also tended to react positively to the microwave oven, and those who reacted negatively tended to react negatively to it. Researchers concluded that “some people tend to like things, whereas others tend to dislike things, and a more thorough understanding of this tendency will lead to a more thorough understanding of the psychology of attitudes.” (Hepler and Albarracı́n 2013)

What are the cases?

```{r haters-cases, exercise=TRUE}
# Answer: 
```

```{r haters-cases-solution}
# Answer: The cases are 200 randomly sampled men and women.
```

What is (are) the response variable(s) in this study?

```{r haters-response, exercise=TRUE}
# Answer: 
```

```{r haters-response-solution}
# Answer: The response variable is attitude towards a fictional microwave oven.
```

What is (are) the explanatory variable(s) in this study?

```{r haters-explanatory, exercise=TRUE}
# Answer: 
```

```{r haters-explanatory-solution}
# Answer: The explanatory variable is dispositional attitude.
```

```{r haters-question, echo=FALSE}
quiz(
  caption="Haters Gonna Hate",
  question_radio("Does the study employ random sampling?", 
           answer("Yes", correct=TRUE, message="The cases are sampled randomly, recruited online using Amazon’s Mechanical Turk."),
           answer("No", message="The cases are sampled randomly, recruited online using Amazon’s Mechanical Turk.")),
  question_radio("Is this an observational study or an experiment?",
                 answer("Experiement", message="This is an observational study since there is no random assignment to treatments."),
                 answer("Observational Study", correct=TRUE, message="This is an observational study since there is no random assignment to treatments."),
                 allow_retry = TRUE),
  question_radio("Can we establish a causal link between the explanatory and response variables?",
                 answer("Yes", message="We cannot establish a causal link between the explanatory and response variables since the study is observational."),
                 answer("No", message="We cannot establish a causal link between the explanatory and response variables since the study is observational.", correct=TRUE),
                 allow_retry = TRUE),
  question_radio("Can the results of the study be generalized to the population at large?",
                 answer("Yes", message="The results of the study can be generalized to the population at large since the sample is random.", correct=TRUE),
                 answer("No", message="The results of the study can be generalized to the population at large since the sample is random."),
               allow_retry = TRUE)
)
```


## Exercise and mental health. 

A researcher is interested in the effects of exercise on mental health, and they propose the following study: use stratified random sampling to ensure representative proportions of 18-30, 31-40 and 41-55 year-olds from the population. Next, randomly assign half the subjects from each age group to exercise twice a week and instruct the rest not to exercise. Conduct a mental health exam at the beginning and at the end of the study and compare the results.

```{r exercise-q1, echo=FALSE}
question_radio("What type of study is this?",
               answer("Experiement", correct=TRUE),
               answer("Observational Study"),
               allow_retry = TRUE)
```


What are the treatment and control groups in this study?

```{r exercise-treatment, exercise=TRUE}
# Answer: 
```

```{r exercise-treatment-solution}
# Answer: The treatment is exercise twice a week and control is no exercise.
```

Does this study make use of blocking? If so, what is the blocking variable?

```{r exercise-blocking, exercise=TRUE}
# Answer: 
```

```{r exercise-blocking-solution}
# Answer: Yes, the blocking variable is age.
```

Does this study make use of blinding?

```{r exercise-blinding, exercise=TRUE}
# Answer: 
```

```{r exercise-blinding-solution}
# Answer: No, the study is not blinded since the patients will know whether or not they are exercising.
```

Comment on whether the results of the study can be used to establish a causal relationship between exercise and mental health and indicate whether the conclusions can be generalized to the population at large.

```{r exercise-causal, exercise=TRUE}
# Answer: 
```

```{r exercise-causal-solution}
# Answer: Since this is an experiment, we can make a causal statement. Since the sample is random, the causal statement can be generalized to the population at large. However, we should be cautious about making a causal statement because of a possible placebo effect.
```

Suppose you are given the task of determining if this proposed study should get funding. Would you have any reservations about the study proposal?

```{r exercise-fund, exercise=TRUE}
# Answer: 
```

```{r exercise-fund-solution}
# Answer: It would be very difficult, if not impossible, to successfully conduct this study since randomly sampled people cannot be required to participate in a clinical trial.
```


## City council survey. 

A city council has requested a household survey be conducted in a suburban area of their city. The area is broken into many distinct and unique neighborhoods, some including large homes, some with only apartments, and others a diverse mixture of housing structures. For each part below, identify the sampling methods described, and describe the statistical pros and cons of the method in the city’s context.


Randomly sample 200 households from the city.

```{r city-council1, exercise=TRUE}
# Answer: 
```

```{r city-council1-solution}
# Answer: Simple random sampling. This is usually an effective method as it assigns equal probability to each household to be picked.

```


Divide the city into 20 neighborhoods, and sample 10 households from each neighborhood.

```{r city-council2, exercise=TRUE}
# Answer: 
```

```{r city-council2-solution}
# Answer: Stratified sampling. This is an effective method in this setting since neighborhoods are unique and this method allows us to sample from each neighborhood.
```

Divide the city into 20 neighborhoods, randomly sample 3 neighborhoods, and then sample all households from those 3 neighborhoods.

```{r city-council3, exercise=TRUE}
# Answer: 
```

```{r city-council3-solution}
# Answer: Cluster sampling. This is not an effective method in this setting since the resulting sample will not contain households from certain neighborhoods and we are told that some neighborhoods are very different from others.
```

Divide the city into 20 neighborhoods, randomly sample 8 neighborhoods, and then randomly sample 50 households from those neighborhoods.

```{r city-council4, exercise=TRUE}
# Answer: 
```

```{r city-council4-solution}
# Answer: Multi-stage sampling. This method will suffer from the same issue discussed in the previous question.
```

Sample the 200 households closest to the city council offices.

```{r city-council5, exercise=TRUE}
# Answer: 
```

```{r city-council5-solution}
# Answer: Convenience sampling. This is not an effective method since it will result in a biased sample for households that are similar to each other (in the same neighborhood) and the sample will not contain any houses from neighborhoods far from the city council offices.
```

## Submit

```{css echo=FALSE}
@media print {
  .topicsContainer,
  .topicActions,
  .exerciseActions .skip {
    display: none;
  }
  .topics .tutorialTitle,
  .topics .section.level2,
  .topics .section.level3:not(.hide) {
    display: block;
  }
  .topics {
    width: 100%;
  }
  .tutorial-exercise, .tutorial-question {
    page-break-inside: avoid;
  }
  .section.level3.done h3 {
    padding-left: 0;
    background-image: none;
  }
  .topics .showSkip .exerciseActions::before {
    content: "Topic not yet completed...";
    font-style: italic;
  }
}
```

Make sure you have the tutorial open in a brower (preferably Google Chrome). Please press `print page` button below to print the tutorial. Make sure to print it to a `pdf` file and upload it into Moodle.

```{js print2pdf1, context="server"}
    // the following 2 chunks print the completed sections of the tutorial to PDF
    // uses "css/print2pdf.css"
    $(document).on('shiny:inputchanged', function(event) {
      if (event.name === 'print2pdf') {
        window.print();
      }
    });
    ```
    
```{r print2pdf2}
    # button can be placed anywhere in the tutorial
    actionButton("print2pdf", "Print page", style="opacity: .7; color: #000;")
```
